# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = 39
## Solution goes here
len(hhcaps.columns)
```

Output:
```
39
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = 55

## Solution goes here
len(hhcaps.State.unique())
```

Output:
```
55
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = 'TX'
## Solution goes here
hhcaps['State'].value_counts()
```

Output:
```
TX    2518
CA    1319
FL    1069
OH     781
IL     735
MI     597
PA     425
MA     261
OK     256
VA     244
IN     229
MN     197
LA     195
CO     187
NC     174
AR     172
AZ     169
MO     169
IA     163
AL     152
NV     143
NY     141
TN     140
WI     119
KS     117
GA     106
KY     104
UT      98
CT      91
NM      77
NE      75
SC      70
WA      63
WV      61
OR      57
MD      53
MS      48
NJ      47
ID      45
PR      45
SD      33
NH      31
DC      28
MT      27
RI      27
DE      26
WY      26
ME      24
ND      18
HI      16
AK      14
VT      11
GU       4
MP       2
VI       1
Name: State, dtype: int64
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = 'ME'

## Solution goes here
best = hhcaps.groupby("State")["Star Rating for health team communicated well with them"].mean().sort_values(ascending=False)
best

```

Output:
```
State
ME    4.650000
WV    4.647059
MS    4.636364
AL    4.507692
LA    4.487805
KY    4.484536
AR    4.384615
TN    4.372881
MT    4.333333
SC    4.333333
MO    4.300885
VT    4.272727
ND    4.250000
NE    4.250000
IA    4.204082
SD    4.166667
KS    4.166667
PA    4.148325
NC    4.138158
NH    4.136364
OK    4.128655
WI    4.101449
IN    4.030075
VI    4.000000
ID    3.972222
RI    3.947368
GA    3.917526
VA    3.865772
DE    3.846154
WY    3.833333
NJ    3.833333
HI    3.818182
MA    3.817308
TX    3.811648
CO    3.770833
MI    3.716279
NM    3.714286
MD    3.695652
IL    3.667820
OH    3.666667
FL    3.645228
WA    3.642857
UT    3.603774
CT    3.550725
MN    3.544554
AK    3.500000
OR    3.409091
NY    3.378641
AZ    3.277108
CA    3.104508
NV    3.000000
PR    2.846154
DC    2.000000
GU         NaN
MP         NaN
Name: Star Rating for health team communicated well with them, dtype: float64
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
answer = 3.6925207756232687

## Solution goes here
Rating = hhcaps.groupby("Type of Ownership")["Star Rating for how patients rated overall care from agency"].mean()
Rating['Hospital Based Program']

```

Output:
```
3.6925207756232687
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
answer = 'Offers Nursing Care Services'

## Solution goes here
hhcaps.nunique()

```

Output:
```
State                                                                                                                                   55
CMS Certification Number (CCN)*                                                                                                      12000
Provider Name                                                                                                                        10394
Address                                                                                                                              11961
City                                                                                                                                  3008
Zip                                                                                                                                   5244
Phone                                                                                                                                11896
Type of Ownership          
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
answer = 9029

## Solution goes here
Offers_cols = [col for col in hhcaps.columns if 'Offers' in col]
Offers_cols
Total_Rows = hhcaps.loc[(hhcaps['Offers Nursing Care Services'] == True) & (hhcaps['Offers Physical Therapy Services'] == True)
                   & (hhcaps['Offers Occupational Therapy Services'] == True) & (hhcaps['Offers Speech Pathology Services'] == True)
                   & (hhcaps['Offers Medical Social Services'] == True) & (hhcaps['Offers Home Health Aide Services'] == True)]
Total_Rows
```

Output:
```
State  CMS Certification Number (CCN)*  \
0        AL                            17000   
1        AL                            17008   
2        AL                            17009   
4        AL                            17014   
6        AL                            17017   
...     ...                              ...   
11991    TX                           747974   
11992    TX                           747975   
11993    TX                           747976   
11998    TX                           747981   
11999    TX                           747982   

                                           Provider Name  \
0      BUREAU OF HOME & COMMUNITY SERVICES           ...   
1      JEFFERSON COUNTY HOME CARE                    ...   
2      ALACARE HOME HEALTH & HOSPICE                 ...   
4      AMEDISYS HOME HEALTH                          ...   
6      SOUTHERN RURAL HEALTH CARE                    ...   
...                                                  ...   
11991  GE
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = 84.0

## Solution goes here
med = hhcaps.loc[hhcaps['Number of completed Surveys'] != 'Not Available']
med["Number of completed Surveys"].agg(['median'])
```

Output:
```
median    84.0
Name: Number of completed Surveys, dtype: float64
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer ==84)

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = 23

## Solution goes here
providers2 = hhcaps[(hhcaps.State == "MO")]
c=0
val = providers2["City"]
for i in val:
    provider4 = i.strip()
    if provider4=='SAINT LOUIS':
        c=c+1
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
answer = 'MS'
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

## Solution goes here
state_P = hhcaps[['State', field]]
State_P2 = state_P[state_P[field] != "Not Available"]
State_P2['Percentage']=State_P2['Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'].astype(float)
State_P2.groupby('State')['Percentage'].mean().sort_values(ascending=False)

```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:7: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: http://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  import sys
State
MS    89.318182
KY    88.627451
WV    88.600000
AL    88.426573
LA    88.287293
AR    87.860870
ME    87.857143
SC    87.365079
MO    86.622222
NC    86.417178
TN    86.241935
KS    86.023529
VI    86.000000
IA    85.951220
TX    85.653549
OK    85.606796
HI    85.583333
PA    85.464286
GA    85.020000
WY    84.857143
VT    84.818182
MA    84.698529
ID    84.658537
IN    84.582353
FL    84.524011
SD    84.440000
NH    84.240000
ND    84.230769
MP    84.000000
AK    83.888889
OH    83.668354
NM    83.580645
RI    83.434783
PR    83.424242
IL    83.358634
VA    83.312500
MT    83.291667
NJ    83.232558
CO    83.158333
WI
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```